# Assignment 3
DATA101: Introduction to Data Science

## Important Dates

 - Released: 02/10/2022
 - Deadline: 02/18/2022
 
### what do you need to submit to OAKS?
* < yourName >A3.ipynb 

Post any questions or any technical issues on Slack.

<div class="alert alert-block alert-success" size = "6">
    <b>Problem 1 </b>
</div>

**Task1**  

Trace the output of the following program. Then rewrite the program using a “while” loop instead of a “for” loop.

```python
langs = ["Python", "Java", "C", "PHP", "Java Script", "Ruby", "Go"]
for i in range(0,len(langs) , 2):
    print (i, langs[i])
```

Output:
```
0 Python
2 C
4 Java Script
6 Go
```

```python
## add your solution here
i = 0
while i < len(langs):
    print(i, langs[i])
    i += 2
```

Output:
```
0 Python
2 C
4 Java Script
6 Go
```

**Task2**

Use list comprehension to create list `num_list` for the following for loop

```python
num = 123450
num_list = []
for i in str(num):
    num_list.append(int(i))
print(num_list)
```

Output:
```
[1, 2, 3, 4, 5, 0]
```

```python
## add your solution here
num = 123450
num_list = []
[num_list.append(int(i)) for i in str(num)]
print(num_list)
```

Output:
```
[1, 2, 3, 4, 5, 0]
```

**Task3**

Create a list `str_list` from `num_list`,  so that each item of `str_list` is the string version of the 5 times of the number in `num_list`.
* For example, [2, 3, 7] would produce ["10", "15", "35"]
  

```python
num_list = [2, 3, 7 ]
## add your solution here
str_list = []
for i in range(len(num_list)):
               str_list.append(str(num_list[i]*5))
print(str_list)
```

Output:
```
['10', '15', '35']
```

<div class="alert alert-block alert-success">
    <b>Problem 2</b>
</div>

**Task1**
* create a Python dictionary called `favorite` with  your favorite restaurants  as keys, and a list of foods/dishes you like in that restaurant as value associate with that key. You should add at leasts two restaurants in your dictionary
* print out all items using the .items() method


```python
## add your code here
favorite = {'Brown Dog Deli':1, 'Rio Chico':2}
favorite['Brown Dog Deli'] = ['salad','fries','tea']
favorite['Rio Chico'] = ['tacos','rice','queso']
for key, value in favorite.items():
    print(key,value)
```

Output:
```
Brown Dog Deli ['salad', 'fries', 'tea']
Rio Chico ['tacos', 'rice', 'queso']
```

**Task2**
* add one more restraunt and food list to `favorite` 


```python
## add your solution here
favorite['Chick-fil-a'] = ['chicken sandwich', 'fries', 'milkshake']
```

**Task3**
* get the food list in `favorite` which associate with ONE restaruant you CHOOSE.  
* print out all foods in that resaruant using the enumerate() method

```python
## add your solution here
enumerateCFA = enumerate(favorite['Chick-fil-a'])
for item in enumerateCFA:
    print(item)
```

Output:
```
(0, 'chicken sandwich')
(1, 'fries')
(2, 'milkshake')
```

**Task4**
* add one more food to ONE restaruant you CHOOSE.  
* print out all foods in that resaruant  using the enumerate() method

```python
## add your solution here
favorite['Chick-fil-a'].append('chicken nuggets')
enumerateCFA = enumerate(favorite['Chick-fil-a'])
for item in enumerateCFA:
    print(item)
```

Output:
```
(0, 'chicken sandwich')
(1, 'fries')
(2, 'milkshake')
(3, 'chicken nuggets')
```

**Task5** (Charllenge one)
* print the restaurants and food lists in an elegant way as follows:
```
    Old Li's Restaurant:
        1. squirrel fish
        2. spicy boiled beef
        3. Dongpo pork
        4. Peking duck
    East Bay Deli:
        1. Chicken Salad Croissant
    CO:
        1. Pad Thai
        2. Pho
 ```   
* hint: use nested loops

```python
## add your solution here
for item in favorite.keys():
    print(item+":")
    for i in range(len(favorite[item])):
        print("\t"+ str(i+1) + ". " +favorite[item][i])
```

Output:
```
Brown Dog Deli:
	1. salad
	2. fries
	3. tea
Rio Chico:
	1. tacos
	2. rice
	3. queso
Chick-fil-a:
	1. chicken sandwich
	2. fries
	3. milkshake
	4. chicken nuggets
```

<div class="alert alert-block alert-success">
    <b>Problem 3</b>
</div>

In this question, you're given two lists: keys and values. Write a loop that uses the `zip` function in order to create a dictionary which contains the specific keys and values.

For example, the function `lists_to_dict( ["one", "two", "three"], [1, 2, 3] )` should return a dictionary that looks like:

`{"one": 1, "two": 2, "three": 3}`

It doesn't matter if, when printed, the key/values are in a different order, so long as the correct keys are paired with the correct values.

**Sorry, using `dict()` is not allowed.** It's a great function that would answer this question in one line, but you can't use it here.

<!-- grade_id: q1b, points: 0, grade: false, solution: true -->
```python
def lists_to_dict(keys, values):
    new_dict = {}
    
    ### BEGIN SOLUTION
    num_list = zip(keys,values)
    for i in range(0, len(keys)):
        new_dict[keys[i]] = values[i]
    ### END SOLUTION
    
    return new_dict
lists_to_dict( ["one", "two", "three"], [1, 2, 3] ) 
```

Output:
```
{'one': 1, 'two': 2, 'three': 3}
```

<!-- grade_id: q1b_test1, points: 10, grade: true, solution: false -->
```python
## for tesing purpose
k1 = ["1", "2", "3"]
v1 = ["one", "two", "three"]
a1 = set({(k1[0], v1[0]), (k1[1], v1[1]), (k1[2], v1[2])})
assert a1 == set(lists_to_dict(k1, v1).items())
```

<!-- grade_id: q1b_test2, points: 10, grade: true, solution: false -->
```python
## for tesing purpose
k2 = ["wat1", "wat2", "wat3", "wat4", "wat5"]
v2 = [1, 2, 3, 4, 5]
a2 = set({(k2[0], v2[0]), (k2[1], v2[1]), (k2[2], v2[2]), (k2[3], v2[3]), (k2[4], v2[4])})
assert a2 == set(lists_to_dict(k2, v2).items())
```

<div class="alert alert-block alert-success">
    <b>Problem 4: Pracice python functions</b>
</div>

**Task1**
Write a function:

 - named `add_one`
 - takes 1 argument, a list
 - returns 1 value, a list
 
In this function, you should loop through each element of the argument list, add 1 to it, and then return the new list.

For example, `add_one([1, 2, 3])` should return `[2, 3, 4]`. No imports!

*HINT*: I would highly recommend constructing a *new* list, rather than updating the *old* one.

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
# add your solution here
def add_one(num_list):
    rtn_list = []
    [rtn_list.append(num+1) for num in num_list]
    return rtn_list
add_one([1,2,3])
```

Output:
```
[2, 3, 4]
```

<!-- grade_id: q1a_test1, points: 5, grade: true, solution: false -->
```python
#for testing purpose
import numpy as np
np.random.seed(7584)

list1 = np.random.randint(-100, 100, 10)
actual1 = list1 + 1
pred1 = add_one(list1.tolist())
assert pred1 == actual1.tolist()
```

**Task2**

Write a function `adjust_wages_by_rate()`, that accepts a list of employee wages and a number rate. The function should modify each wage in the list by multipling the rate, and return a list of adjusted wages
* For example:
`wages` [20.00, 30.00, 15.00, 19.50, 13.00] and `rate` 1.5 will produce [30.0, 45.0, 22.5, 29.25, 19.5]
* write statements to test the function using above example



```python
# add your solution here
def adjust_wages_by_rate(wages_list,rate):
    rtn_list = []
    [rtn_list.append(wage*rate) for wage in wages_list]
    return rtn_list
wages = [20.00, 30.00, 15.00, 19.50, 13.00]
adjust_wages_by_rate(wages,1.5)
```

Output:
```
[30.0, 45.0, 22.5, 29.25, 19.5]
```

**Task3**

Write a function adjust_wages_by_rates(), that accepts a list of employee wages and a list of corresponding rates for each employee. The function should modify each wage in the list by  multipling the corresponding rate, and return a list of adjusted wages.(Surggestion: using zip())

* for example: wages [20.00, 30.00, 15.00, 19.50, 13.00] and rates [1.5, 0.9, 1.2, 1.0, 2] will produce [30.0, 27.0, 18.0, 19.5, 26.0]

* write statements to test the function using above example



```python
# add your solution here
def adjust_wages_by_rates(wages,rates):
    rtn_list = []
    [rtn_list.append(item1 * item2) for item1, item2 in zip(wages, rates)]
    return rtn_list
adjust_wages_by_rates( [20.00, 30.00, 15.00, 19.50, 13.00],[1.5, 0.9, 1.2, 1.0, 2])
```

Output:
```
[30.0, 27.0, 18.0, 19.5, 26.0]
```

<div class="alert alert-block alert-success">
    <b>Problem 4: Finding all</b>
</div>

**Task1** 
Write a function that finds all occurrences of a specific value in a list, and records--in a separate list--these occurrences in terms of the *indices* where they were found.

Your function should:

 - be named `find_all`
 - take 2 arguments: a list of items, and a single element to search for in the list
 - returns 1 list: a list of *indices* into the input list that correspond to elements in the input list that match what we were looking for
 
For example, `find_all([1, 2, 3, 4, 5, 2], 2)` would return `[1, 5]`. `find_all([1, 2, 3], 0)` should return `[]`.

*HINT*: Instead of doing a simple `for element in list:` loop, try using the `range()` function in the loop header.

You cannot use any built-in functions besides `len()` and `range()`.

```python
## add your solution here..
def find_all(item_list, search):
    rtn_list = []
    for i in range(len(item_list)):
        if item_list[i] == search:
            rtn_list.append(i)
    return rtn_list
find_all([1,2,3,4,5,2],2)
```

Output:
```
[1, 5]
```

<!-- grade_id: q3b_test1, points: 5, grade: true, solution: false -->
```python
##for testing purpose
l1 = [1, 2, 3, 4, 5, 2]
s1 = 2
a1 = [1, 5]
assert set(a1) == set(find_all1(l1, s1))
```

Output:
```
NameError                                 Traceback (most recent call last)
/var/folders/vm/39l8g9l10r5cmk_8kpg_b4sm0000gn/T/ipykernel_26443/2863872079.py in <module>
      3 s1 = 2
      4 a1 = [1, 5]
----> 5 assert set(a1) == set(find_all1(l1, s1))

NameError: name 'find_all1' is not defined
```

<!-- grade_id: q3b_test2, points: 5, grade: true, solution: false -->
```python
#for testing purpose
l2 = ["a", "random", "set", "of", "strings", "for", "an", "interesting", "strings", "problem"]
s2 = "strings"
a2 = [4, 8]
assert set(a2) == set(find_all(l2, s2))
```

<!-- grade_id: q3b_test3, points: 5, grade: true, solution: false -->
```python
#for testing purpose
l3 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
s3 = 1001
a3 = []
assert set(a3) == set(find_all(l3, s3))
```

**Task2**

Write a function which provides counts of very specific elements in a list. This function will return a dictionary where the keys are the elements you wanted to find, and their values are the number of times (counts) they were found in the target list.

*HINT*: You can use your answer to Task1 here to expedite the solution! You don't have to, but you can, as it partially solves the problem here. To do that, simply call your function `find_all`--you do NOT need to copy/paste it into the cell below! Python will see it in an earlier cell and will know that it exists (as long as you've hit the "Play/Run" button on that cell, of course. And if you make any changes to it, you should hit the "Play/Run" button again; this updates it in the Python program, kind of like hitting "Save" on a Word doc).

Your function should:

 - be named `element_counts`
 - take 2 arguments, both lists: a list of your data, and a list of elements you want counted in your data
 - return a dictionary: keys are the elements you wanted counted, and values are their counts in the data
 
For example, `element_counts([1, 2, 3, 4, 5, 2], [2, 5])` would return `{2: 2, 5: 1}`, as there were two 2s in the data list, and one 5. Another example would be `element_counts([1, 2, 3], [0, 1])`, which would return `{0: 0, 1: 1}`, as there were no 0s in the original list, and one 1.

You cannot use any built-in functions besides `len()` and `range()`.

```python
# add your solution here
def element_counts(data_list,element_list):
    rtnDict = {}
    for i in range(len(element_list)):
        rtnDict[element_list[i]] = len(find_all(data_list,element_list[i]))
    return rtnDict
element_counts([1, 2, 3, 4, 5, 2], [2, 5])
```

Output:
```
{2: 2, 5: 1}
```

<!-- grade_id: q3c_test1, points: 5, grade: true, solution: false -->
```python
##for testing purpose
l1 = [1, 2, 3, 4, 5, 2]
s1 = [2, 5]
a1 = {2: 2, 5: 1}
assert a1 == element_counts3(l1, s1)
```

Output:
```
NameError                                 Traceback (most recent call last)
/var/folders/vm/39l8g9l10r5cmk_8kpg_b4sm0000gn/T/ipykernel_26443/135847741.py in <module>
      3 s1 = [2, 5]
      4 a1 = {2: 2, 5: 1}
----> 5 assert a1 == element_counts3(l1, s1)

NameError: name 'element_counts3' is not defined
```

<!-- grade_id: q3c_test2, points: 5, grade: true, solution: false -->
```python
#for testing purpose
l2 = ["a", "random", "set", "of", "strings", "for", "an", "interesting", "strings", "problem"]
s2 = ["strings", "of", "notinthelist"]
a2 = {"strings": 2, "of": 1, "notinthelist": 0}
assert a2 == element_counts(l2, s2)
```